---
title: 'Chapter 1: R Basics'
description: ""
---

## Insert exercise title here

```yaml
type: NormalExercise 
xp: 100 
key: 2fff329b50   
```





`@instructions`


`@hint`


`@pre_exercise_code`

```{r}

```


`@sample_code`

```{r}

```


`@solution`

```{r}

```


`@sct`

```{r}

```


`@possible_answers`


`@feedback`


