---
title: 'Reshaping data'
description: ""
---

## Insert exercise title here

```yaml
type: NormalExercise 
xp: 100 
key: 44e4b50e28   
```


sdasad


`@instructions`


`@hint`


`@pre_exercise_code`

```{r}
library('dplyr')
library('reshape2')
cars<-mtcars
```


`@sample_code`

```{r}

```


`@solution`

```{r}

```


`@sct`

```{r}

```


