---
title: 'Reshaping data'
description: ""
---

## Insert exercise title here

```yaml
type: VideoExercise 
xp: 50 
key: 813d0fea50   
```

`@projector_key`
388773edc7a9e64c176bf6595592a2c5
---

## Insert exercise title here

```yaml
type: NormalExercise 
xp: 100 
key: 44e4b50e28   
```


sdasad


`@instructions`


`@hint`


`@pre_exercise_code`

```{r}
library('reshape2')
library('dplyr')
#library('data.table')
cars<- mtcars
```


`@sample_code`

```{r}

```


`@solution`

```{r}

```


`@sct`

```{r}

```


