---
title: 'Chapter 1: R Basics'
description: ""
---

## R Basics / Entering commands

```yaml
type: NormalExercise 
xp: 100 
key: 2fff329b50   
```


R is a command line driven program. It means that you enter your commands at a prompt (R Console), and an output is produced after you press "Enter".

Normally, people write their codes in a "Script.R" file, which is equivalent to a text file. That allows the user to save their code to work on it later on.  

In you right hand side you see both the _Sript.R_ file, and the _R Console_. You can write your code in the _Srcipt.R_ window and see the output of your code in the _R Console_ by pressing the button **Run Code.**


`@instructions`
Please write the following command in the Script.R window and see the output in your R Console

`` 2+2 ``

Once you see the output press **"Submit Answer"**

`@hint`


`@pre_exercise_code`

```{r}

```


`@sample_code`

```{r}
# Write your code here
```


`@solution`

```{r}
2+2
```


`@sct`

```{r}

```


`@possible_answers`


`@feedback`


